# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

```julia
length(myfriends)
```

Output:
```
5
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

```julia
for n = 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
for n ∈ 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia
m, n = 5, 5
A = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

Here's some syntactic sugar for the same nested `for` loop

```julia
B = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

```julia
A == B
```

Output:
```
true
```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia
for i ∈ 1:100
    println("$(i)^2 \t = \t $(i^2)")
end
```

Output:
```
1^2 	 = 	 1
2^2 	 = 	 4
3^2 	 = 	 9
4^2 	 = 	 16
5^2 	 = 	 25
6^2 	 = 	 36
7^2 	 = 	 49
8^2 	 = 	 64
9^2 	 = 	 81
10^2 	 = 	 100
11^2 	 = 	 121
12^2 	 = 	 144
13^2 	 = 	 169
14^2 	 = 	 196
15^2 	 = 	 225
16^2 	 = 	 256
17^2 	 = 	 289
18^2 	 = 	 324
19^2 	 = 	 361
20^2 	 = 	 400
21^2 	 = 	 441
22^2 	 = 	 484
23^2 	 = 	 529
24^2 	 = 	 576
25^2 	 = 	 625
26^2 	 = 	 676
27^2 	 = 	 729
28^2 	 = 	 784
29^2 	 = 	 841
30^2 	 = 	 900
31^2 	 = 	 961
32^2 	 = 	 1024
33^2 	 = 	 1089
34^2 	 = 	 1156
35^2 	 = 	 1225
36^2 	 = 	 1296
37^2 	 = 	 1369
38^2 	 = 	 1444
39^2 	 = 	 1521
40^2 	 = 	 1600
41^2 	 = 	 1681
42^2 	 = 	 1764
43^2 	 = 	 1849
44^2 	 = 	 1936
45^2 	 = 	 2025
46^2 	 = 	 2116
47^2 	 = 	 2209
48^2 	 = 	 2304
49^2 	 = 	 2401
50^2 	 = 	 2500
51^2 	 = 	 2601
52^2 	 = 	 2704
53^2 	 = 	 2809
54^2 	 = 	 2916
55^2 	 = 	 3025
56^2 	 = 	 3136
57^2 	 = 	 3249
58^2 	 = 	 3364
59^2 	 = 	 3481
60^2 	 = 	 3600
61^2 	 = 	 3721
62^2 	 = 	 3844
63^2 	 = 	 3969
64^2 	 = 	 4096
65^2 	 = 	 4225
66^2 	 = 	 4
```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia
squares = Dict{Int, Int}()
for i ∈ 1:100
    squares[i] = i^2
end
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
# squares_arr = [i^2 for i ∈ 1:100]
squares_arr = [i^2 for i in 1:100];
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

Please click on `Validate` on the top, once you are done with the exercises.